# Assignment 1 Part 2
## Grades

The dataset file in [assets/grades.txt](assets/grades.txt) contains multiple lines of people along with their grades in a class. Fix the incorrect regex between `### BEGIN CODE` and `### END CODE` to generate a list of just those students who received a B in the course (e.g., \[\'John Doe\', \'Jane Doe\'\].)

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re

def student_grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    ### BEGIN CODE
    pattern = "[a-zA-Z]* [a-zA-Z]*: B"
    matches = re.findall(pattern,grades)
    for i in range(len(matches)):
        matches[i] = matches[i][:-3]
    
    ### END CODE
        
        
    # YOUR CODE HERE
    
    #raise NotImplementedError()
    

    return matches  
    
student_grades()

```

Output:
```
['Bell Kassulke',
 'Simon Loidl',
 'Elias Jovanovic',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Yasmin Dar',
 'Viswamitra Upandhye',
 'Killian Kaufman',
 'Elwood Page',
 'Elodie Booker',
 'Adnan Chen',
 'Hank Spinka',
 'Hannah Bayer']
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(student_grades()) == 16
```