# Assignment 2 - Part 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 2
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.00779`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    import pandas as pd
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)
    #P_NUMVRC - total number of varicella doses, look for 1 or greater
    #HAD_CPOX - value will be 1 if they did have it
    #SEX - male 1 or female 2
    #males = df[df['P_NUMVRC'].gt(0) & df['HAD_CPOX'].eq(1) & df['SEX'].eq(1)]
    #females = df[df['P_NUMVRC'].gt(0) & df['HAD_CPOX'].eq(1) & df['SEX'].eq(2)]

    yesChickenPox = df[df['P_NUMVRC'].gt(0) & df['HAD_CPOX'].eq(1)]
    noChickenPox = df[df['P_NUMVRC'].gt(0) & df['HAD_CPOX'].eq(2)]

    malesYes = yesChickenPox[yesChickenPox['SEX'].eq(1)]
    malesNo = noChickenPox[noChickenPox['SEX'].eq(1)]
    femalesYes = yesChickenPox[yesChickenPox['SEX'].eq(2)]
    femalesNo = noChickenPox[noChickenPox['SEX'].eq(2)]

    results = {}
    results['male'] = len(malesYes)/len(malesNo)
    results['female'] = len(femalesYes)/len(femalesNo)

    return results    
    
    #raise NotImplementedError()
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

```python

```